# CSCI-A201/597 Introduction to Programming I: 
# Final Exam: Programming Part

* Answer each questions as instructed. 

* When you finish the exam:
    * Step 1: Submit the print-out to intructor/TA.
    * Step 2: Save the file.
    * Step 3: Download the file `.ipynb`:
    * Step 4: Submission to Autograder: https://iu.instructure.com/courses/2284109/assignments/17304412. 

## p1.(1).  [10 pts]

```python
# finish the program: 
p1_1 = None
```

```python
print(p1_1)
```

```python

```

## p1.(2).  [10 pts]

```python
# finish the function:

def p1_2(pair_list):
    pass
```

```python
# test 1
p1_2(['apple:27', 'dog:2','apple:1'])
```

```python
# test 2:
p1_2(['cat:5', 'bird:3', 'fish:7'])
```

```python

```

## p1.(3).  [10 pts]

```python
# finish the function:

def p1_3(set1, set2):
    pass
```

```python
# test 1:
p1_3({"bug", "rice", "apple"}, {"rice", "bug", "sugar"})
```

```python
# test 2:
p1_3({"apple", "banana"}, {"carrot", "beet"})
```

```python
# test 3:
p1_3({"apple", "banana"}, {"apple", "beet"})
```

```python

```

## p1.(4)  [10 pts]

```python
# finish the function:

def p1_4(*args):
    pass
```

```python
# test 1
p1_4()
```

```python
# test 2
p1_4(1, 2)
```

```python
# test 3
p1_4("happy", [1, 2], 3)
```

```python

```

### P2.(1) [10 pts]

```python
# P2 (1)

def p2_1(input_dict):
    pass

```

```python
# test
d = {1:10, 2:20, 3:30}
p2_1(d)
```

```python
# test
d = {1:10, 2:20, 3:30, 4:30}
p2_1(d)
```

```python
# test
d = {4:True, 2:True, 0:True}
p2_1(d) 
```

```python

```

## p2.(2)  [10 pts]

```python
# write your code here:

def p2_2(data):
    pass
```

```python
# test 1
data1 = {"Math": {"Alice": "A", "Bob": "A", "Charlie": "A"}}
p2_2(data1)
```

```python
# test 2
data2 = {"Math": {"Alice": "A", "Bob": "B", "Charlie": "A"},
    "Science": {"Alice": "A", "Bob": "A", "Charlie": "A"}}
p2_2(data2)
```

```python
# test 3
data3 = { "Math": {"Alice": "A", "Bob": "B", "Charlie": "A"},
    "Science": {"Alice": "A", "Bob": "A", "Charlie": "B"},
    "History": {"Alice": "B", "Bob": "A", "Charlie": "A"}}
p2_2(data3)
```

## p2.(3)  [10 pts]

```python
# write your code here:

def p2_3(data):
    pass

```

```python
# test 1

data4 = {
    "Math": {"Alice": "A", "Bob": "B", "Eva": "B", "Frank": "B-"}, 
    "Science": {"Alice": "B", "Bob": "A", "Eva": "Z", "Frank": "B"}, 
    "History": {"Alice": "C", "Bob": "B", "Eva": "A", "Frank": "C"}}
p2_3(data4)
```

```python
# test 2

data5 = {"Math": {"Alice": "A"}, "Science": {"Alice": "B"}}
p2_3(data5) 
```

```python
# test 3

data6 = {"Math": {"Alice": "A-"}, "Science": {"Alice": "B+"}}
p2_3(data6)
```

```python

```

## p2.(4)  [10 pts]

```python
# write your code here:

def p2_4(data, filename):
    # Simplified grading standard
    grade_points = {"A": 4.0, "B": 3.0, "C": 2.0, "D": 1.0, "F": 0.0}
    
    # finish the program below

    

```

```python
# test:

data = {
    "Math": {"Alice": "A", "Bob": "B", "Charlie": "A", "David": "A+"},
    "History": {"Alice": "A", "Bob": "A"},
    "Science": {"Alice": "A-", "Bob": "B+", "Charlie": "Z"}}

p2_4(data, "average_grades.txt")
```

```python

```

## p3  [10 pts]

```python
# write your code here:

def p3(user_date):
    # modify the program:
    months = {
    "January": 1, "February": 2, "March": 3, "April": 4,
    "May": 5, "June": 6, "July": 7, "August": 8,
    "September": 9, "October": 10, "November": 11, "December": 12}
    yyyy = 0
    mm = 0
    dd = 0
    return f"{yyyy:04}-{mm:02}-{dd:02}" # # Return ISO format
    
    
```

```python
# test 1
p3("March 15, 2025")
```

```python
# test 2
p3("March 100, 2025")
```

```python
# test 3
p3("12/30/2024")
```

```python
# test 4
p3("15/3/2024")
```

```python
# test 5
p3("march 15, 2025")
```

```python
# test 6
p3("15 March, 2023")
```

```python

```

Tests for p1.(1) (points: 10)
```
>>> def check_p1_1():
...     assert isinstance(p1_1, dict), 'not a dictionary'
...     assert len(p1_1) == 50
...     for i in range(2, 101, 2):
...         assert i in p1_1, '{i} should be included'
...         assert i * i == p1_1[i]
>>> check_p1_1()
```


Tests for p1.(2) (points: 10)
```
>>> pair_list = ['apple:27', 'dog:2', 'apple:1']
>>> result = p1_2(pair_list)
>>> expected = {'apple': 28, 'dog': 2}
>>> assert result == expected, f'Test Case 1 Failed: Expected {expected}, got {result}'
```
```
>>> pair_list = ['cat:5', 'bird:3', 'fish:7']
>>> result = p1_2(pair_list)
>>> expected = {'cat': 5, 'bird': 3, 'fish': 7}
>>> assert result == expected, f'Test Case 2 Failed: Expected {expected}, got {result}'
```
```
>>> pair_list = ['lion:9']
>>> result = p1_2(pair_list)
>>> expected = {'lion': 9}
>>> assert result == expected, f'Test Case 3 Failed: Expected {expected}, got {result}'
```
```
>>> pair_list = []
>>> result = p1_2(pair_list)
>>> expected = {}
>>> assert result == expected, f'Test Case 4 Failed: Expected {expected}, got {result}'
```
```
>>> pair_list = ['banana:10', 'banana:20', 'banana:30']
>>> result = p1_2(pair_list)
>>> expected = {'banana': 60}
>>> assert result == expected, f'Test Case 5 Failed: Expected {expected}, got {result}'
```
```
>>> pair_list = ['grape:0', 'grape:5']
>>> result = p1_2(pair_list)
>>> expected = {'grape': 5}
>>> assert result == expected, f'Test Case 6 Failed: Expected {expected}, got {result}'
```
```
>>> pair_list = ['key1:10', 'key1:-3', 'key2:-7']
>>> result = p1_2(pair_list)
>>> expected = {'key1': 7, 'key2': -7}
>>> assert result == expected, f'Test Case 7 Failed: Expected {expected}, got {result}'
```
```
>>> pair_list = ['a:1000000', 'a:2000000', 'b:3000000']
>>> result = p1_2(pair_list)
>>> expected = {'a': 3000000, 'b': 3000000}
>>> assert result == expected, f'Test Case 8 Failed: Expected {expected}, got {result}'
```


Tests for p1.(3) (points: 10)
```
>>> set1 = {'cat', 'dog', 'bug'}
>>> set2 = {'dog', 'rabbit', 'bug'}
>>> result = p1_3(set1, set2)
>>> assert result == {'dog'}, f"Test Case 1 Failed: Expected {{'dog'}}, got {result}"
```
```
>>> set1 = {'apple', 'banana'}
>>> set2 = {'carrot', 'beet'}
>>> result = p1_3(set1, set2)
>>> assert result == set(), f'Test Case 2 Failed: Expected set(), got {result}'
```
```
>>> set1 = {'bug'}
>>> set2 = {'bug'}
>>> result = p1_3(set1, set2)
>>> assert result == set(), f'Test Case 3 Failed: Expected set(), got {result}'
```
```
>>> set1 = {'apple', 'banana', 'cherry'}
>>> set2 = {'banana', 'cherry', 'date'}
>>> result = p1_3(set1, set2)
>>> assert result == {'banana', 'cherry'}, f"Test Case 4 Failed: Expected {{'banana', 'cherry'}}, got {result}"
```
```
>>> set1 = set()
>>> set2 = set()
>>> result = p1_3(set1, set2)
>>> assert result == set(), f'Test Case 5 Failed: Expected set(), got {result}'
```
```
>>> set1 = {'apple', 'bug'}
>>> set2 = set()
>>> result = p1_3(set1, set2)
>>> assert result == set(), f'Test Case 6 Failed: Expected set(), got {result}'
```
```
>>> set1 = {'apple', 'bug'}
>>> set2 = {'apple', 'banana'}
>>> result = p1_3(set1, set2)
>>> assert result == {'apple'}, f"Test Case 7 Failed: Expected {{'apple'}}, got {result}"
```
```
>>> set1 = {'bug', 'carrot', 'date'}
>>> set2 = {'bug', 'carrot', 'date'}
>>> result = p1_3(set1, set2)
>>> assert result == {'carrot', 'date'}, f"Test Case 8 Failed: Expected {{'carrot', 'date'}}, got {result}"
```


Tests for p1.(4) (points: 10)
```
>>> result = p1_4()
>>> expected = 0
>>> assert result == expected, f'Test Case 1 Failed: Expected {expected}, got {result}'
```
```
>>> result = p1_4(10)
>>> expected = 1
>>> assert result == expected, f'Test Case 2 Failed: Expected {expected}, got {result}'
```
```
>>> result = p1_4(1, 2, 3, 4, 5)
>>> expected = 5
>>> assert result == expected, f'Test Case 3 Failed: Expected {expected}, got {result}'
```
```
>>> result = p1_4('apple', 42, [1, 2, 3], {'key': 'value'}, 3.14)
>>> expected = 5
>>> assert result == expected, f'Test Case 4 Failed: Expected {expected}, got {result}'
```
```
>>> result = p1_4([1, 2, 3], (4, 5), {6, 7}, {'key': 'value'})
>>> expected = 4
>>> assert result == expected, f'Test Case 5 Failed: Expected {expected}, got {result}'
```
```
>>> large_args = range(1000)
>>> result = p1_4(*large_args)
>>> expected = 1000
>>> assert result == expected, f'Test Case 6 Failed: Expected {expected}, got {result}'
```


Tests for p2.(1) (points: 10)
```
>>> assert p2_1({1: 10, 2: 20, 3: 30}) == {10: [1], 20: [2], 30: [3]}, 'Test Case 1 Failed'
```
```
>>> assert p2_1({1: 10, 2: 20, 3: 30, 4: 30}) == {10: [1], 20: [2], 30: [3, 4]}, 'Test Case 2 Failed'
```
```
>>> assert p2_1({4: True, 2: True, 0: True}) == {True: [0, 2, 4]}, 'Test Case 3 Failed'
```
```
>>> assert p2_1({'a': 1, 'b': 2, 'c': 1}) == {1: ['a', 'c'], 2: ['b']}, 'Test Case 4 Failed'
```
```
>>> assert p2_1({}) == {}, 'Test Case 5 Failed'
```
```
>>> assert p2_1({5: 'x', 3: 'x', 7: 'x'}) == {'x': [3, 5, 7]}, 'Test Case 6 Failed'
```
```
>>> assert p2_1({1: 100, 2: 200, 3: 100, 4: 300}) == {100: [1, 3], 200: [2], 300: [4]}, 'Test Case 7 Failed'
```
```
>>> assert p2_1({'alpha': 'a', 'beta': 'b', 'gamma': 'a'}) == {'a': ['alpha', 'gamma'], 'b': ['beta']}, 'Test Case 8 Failed'
```


Tests for p2.(2) (points: 10)
```
>>> data1 = {'Math': {'Alice': 'A', 'Bob': 'A', 'Charlie': 'A'}}
>>> assert p2_2(data1) == {'Alice', 'Bob', 'Charlie'}, 'Test Case 1 Failed'
```
```
>>> data2 = {'Math': {'Alice': 'A', 'Bob': 'B', 'Charlie': 'A'}, 'Science': {'Alice': 'A', 'Bob': 'A', 'Charlie': 'A'}}
>>> assert p2_2(data2) == {'Alice', 'Charlie'}, 'Test Case 2 Failed'
```
```
>>> data3 = {'Math': {'Alice': 'A', 'Bob': 'B', 'Charlie': 'A'}, 'Science': {'Alice': 'A', 'Bob': 'A', 'Charlie': 'B'}, 'History': {'Alice': 'B', 'Bob': 'A', 'Charlie': 'A'}}
>>> assert p2_2(data3) == set(), 'Test Case 3 Failed'
```
```
>>> data4 = {'Math': {'Alice': 'B', 'Bob': 'C', 'Charlie': 'B'}, 'Science': {'Alice': 'B', 'Bob': 'A', 'Charlie': 'C'}, 'History': {'Alice': 'C', 'Bob': 'B', 'Charlie': 'B'}}
>>> assert p2_2(data4) == set(), 'Test Case 4 Failed'
```
```
>>> data5 = {'Math': {'Alice': 'A', 'Bob': 'C', 'Charlie': 'B'}, 'Science': {'Alice': 'B', 'Bob': 'C', 'Charlie': 'A'}, 'History': {'Alice': 'C', 'Bob': 'B', 'Charlie': 'C'}}
>>> assert p2_2(data5) == set(), 'Test Case 5 Failed'
```
```
>>> data6 = {'Math': {'Alice': 'A', 'Bob': 'A'}, 'Science': {'Alice': 'A', 'Bob': 'A'}, 'History': {'Alice': 'A', 'Bob': 'A'}}
>>> assert p2_2(data6) == {'Alice', 'Bob'}, 'Test Case 6 Failed'
```
```
>>> data7 = {'Math': {'Alice': 'A', 'Bob': 'B', 'Charlie': 'A'}, 'Science': {'Alice': 'A', 'Charlie': 'A'}, 'History': {'Alice': 'A', 'Charlie': 'A'}}
>>> assert p2_2(data7) == {'Alice', 'Charlie'}, 'Test Case 7 Failed'
```


Tests for p2.(3) (points: 10)
```
>>> data4 = {'Math': {'Alice': 'A', 'Bob': 'B', 'Eva': 'B', 'Frank': 'B-'}, 'Science': {'Alice': 'B', 'Bob': 'A', 'Eva': 'Z', 'Frank': 'B'}, 'History': {'Alice': 'C', 'Bob': 'B', 'Eva': 'A', 'Frank': 'C'}}
>>> assert p2_3(data4) == {'Math': 3.33, 'Science': 3.33}
```
```
>>> data5 = {'Math': {'Alice': 'A'}, 'Science': {'Alice': 'B'}}
>>> assert p2_3(data5) == {'Math': 4.0}
```
```
>>> data6 = {'Math': {'Alice': 'A-'}, 'Science': {'Alice': 'B+'}}
>>> assert p2_3(data6) == {}
```
```
>>> data7 = {'Math': {'Alice': 'A', 'Bob': 'A', 'Charlie': 'A'}, 'Science': {'Alice': 'B', 'Bob': 'B', 'Charlie': 'B'}, 'History': {'Alice': 'C', 'Bob': 'C', 'Charlie': 'C'}}
>>> assert p2_3(data7) == {'Math': 4.0}, 'Test Case 4 Failed'
```
```
>>> data8 = {'Math': {'Alice': 'A', 'Bob': 'B', 'Charlie': 'X'}, 'Science': {'Alice': 'B', 'Bob': 'A', 'Charlie': 'C'}, 'History': {'Alice': 'C', 'Bob': 'D', 'Charlie': 'C'}}
>>> assert p2_3(data8) == {'Math': 3.5}, 'Test Case 5 Failed'
```
```
>>> data9 = {'Math': {'Alice': 'A'}, 'Science': {'Bob': 'B+'}, 'History': {'Charlie': 'Z'}}
>>> assert p2_3(data9) == {'Math': 4.0}, 'Test Case 6 Failed'
```
```
>>> data10 = {'Math': {'Alice': 'A', 'Bob': 'B', 'Charlie': 'A', 'David': 'B'}, 'Science': {'Alice': 'B', 'Bob': 'A', 'Charlie': 'B', 'David': 'A'}, 'History': {'Alice': 'A', 'Bob': 'B', 'Charlie': 'A', 'David': 'B'}}
>>> assert p2_3(data10) == {'Math': 3.5, 'Science': 3.5, 'History': 3.5}, 'Test Case 7 Failed'
```
```
>>> data11 = {'Math': {'Alice': 'A', 'Bob': 'B', 'Charlie': 'A', 'David': 'C', 'Eva': 'B', 'Frank': 'B-', 'Grace': 'B', 'Hannah': 'A', 'Ivy': 'C', 'Jack': 'D'}, 'Science': {'Alice': 'B', 'Bob': 'A', 'Charlie': 'X', 'David': 'A', 'Eva': 'C', 'Frank': 'B', 'Grace': 'C', 'Hannah': 'B', 'Ivy': 'A', 'Jack': 'B+'}, 'History': {'Alice': 'C', 'Bob': 'B', 'Charlie': 'A', 'David': 'B', 'Eva': 'A', 'Frank': 'C', 'Grace': 'B', 'Hannah': 'B', 'Ivy': 'A', 'Jack': 'F'}, 'Geography': {'Alice': 'A', 'Bob': 'A', 'Charlie': 'B', 'David': 'B', 'Eva': 'A', 'Frank': 'B', 'Grace': 'A', 'Hannah': 'A', 'Ivy': 'B', 'Jack': 'A'}, 'Literature': {'Alice': 'B', 'Bob': 'C', 'Charlie': 'A', 'David': 'C', 'Eva': 'A', 'Frank': 'B', 'Grace': 'D', 'Hannah': 'B', 'Ivy': 'C', 'Jack': 'F'}}
>>> assert p2_3(data11) == {'Geography': 3.6}, 'Test Case 8 Failed'
```


Tests for p2.(4) (points: 10)
```
>>> import os
>>> data1 = {'Math': {'Alice': 'A', 'Bob': 'B', 'Charlie': 'A'}, 'Science': {'Alice': 'A', 'Bob': 'A', 'Charlie': 'A'}}
>>> filename1 = 'test_output1.txt'
>>> p2_4(data1, filename1)
>>> with open(filename1, 'r') as file:
...     content = file.read()
...     assert 'Math:3.67' in content
...     assert 'Science:4' in content
>>> os.remove(filename1)
```
```
>>> import os
>>> data2 = {'Math': {'Alice': 'A', 'Bob': 'Z', 'Charlie': 'B'}, 'Science': {'Alice': 'A', 'Bob': 'B', 'Charlie': 'C'}}
>>> filename2 = 'test_output2.txt'
>>> p2_4(data2, filename2)
>>> with open(filename2, 'r') as file:
...     content = file.read()
...     assert 'Math:3.5' in content
...     assert 'Science:3.0' in content
>>> os.remove(filename2)
```
```
>>> import os
>>> data3 = {}
>>> filename3 = 'test_output3.txt'
>>> p2_4(data3, filename3)
>>> with open(filename3, 'r') as file:
...     content = file.read()
...     assert len(content) < 2
>>> os.remove(filename3)
```
```
>>> import os
>>> data4 = {'History': {'Alice': 'A', 'Bob': 'A', 'Charlie': 'A'}}
>>> filename4 = 'test_output4.txt'
>>> p2_4(data4, filename4)
>>> with open(filename4, 'r') as file:
...     content = file.read()
...     assert 'History:4.0' in content
>>> os.remove(filename4)
```
```
>>> import os
>>> data5 = {'History': {'Alice': 'Z', 'Bob': 'Y'}}
>>> filename5 = 'test_output5.txt'
>>> p2_4(data5, filename5)
>>> with open(filename5, 'r') as file:
...     assert file.read() == ''
>>> os.remove(filename5)
```
```
>>> import os
>>> data6 = {'Math': {'Alice': 'B', 'Bob': 'C', 'Charlie': 'F'}, 'Science': {'Alice': 'B', 'Bob': 'Z', 'Charlie': 'C'}, 'History': {'Alice': 'D', 'Bob': 'D', 'Charlie': 'D'}}
>>> filename6 = 'test_output6.txt'
>>> p2_4(data6, filename6)
>>> with open(filename6, 'r') as file:
...     content = file.read()
...     assert 'Math:1.67' in content
...     assert 'Science:2.5' in content
...     assert 'History:1.0' in content
>>> os.remove(filename6)
```
```
>>> import os
>>> data7 = {'Math': {}, 'Science': {}, 'History': {}}
>>> filename7 = 'test_output7.txt'
>>> p2_4(data7, filename7)
>>> with open(filename7, 'r') as file:
...     assert file.read() == ''
>>> os.remove(filename7)
```


Tests for p3 (points: 10)
```
>>> assert p3('3/15/2023') == '2023-03-15'
```
```
>>> assert p3('12/31/2020') == '2020-12-31'
```
```
>>> assert p3('1/1/1999') == '1999-01-01'
```
```
>>> assert p3('March 15, 2023') == '2023-03-15'
```
```
>>> assert p3('December 31, 2020') == '2020-12-31'
```
```
>>> assert p3('January 1, 1999') == '1999-01-01'
```
```
>>> try:
...     p3('15/2023')
... except ValueError as e:
...     assert str(e) == 'Invalid date format or value'
```
```
>>> try:
...     p3('13/15/2023')
... except ValueError as e:
...     assert str(e) == 'Invalid date format or value'
```
```
>>> try:
...     p3('12/32/2023')
... except ValueError as e:
...     assert str(e) == 'Invalid date format or value'
```
```
>>> try:
...     p3('15 March, 2023')
... except ValueError as e:
...     assert str(e) == 'Invalid date format or value'
```
```
>>> try:
...     p3('Mar 15, 2023')
... except ValueError as e:
...     assert str(e) == 'Invalid date format or value'
```
```
>>> try:
...     p3('March, 2023')
... except ValueError as e:
...     assert str(e) == 'Invalid date format or value'
```
```
>>> assert p3('1/1/0001') == '0001-01-01'
>>> assert p3('January 1, 1') == '0001-01-01'
```
```
>>> assert p3('12/31/9999') == '9999-12-31'
>>> assert p3('December 31, 9999') == '9999-12-31'
```
```
>>> try:
...     p3('3/15/abcd')
... except ValueError as e:
...     assert str(e) == 'Invalid date format or value'
```
```
>>> try:
...     p3('3/15/2023 and March 15, 2023')
... except ValueError as e:
...     assert str(e) == 'Invalid date format or value'
```
```
>>> try:
...     p3('3/March/2023')
... except ValueError as e:
...     assert str(e) == 'Invalid date format or value'
```
