# CSCI-A202/598 Introduction to Programming I: 
# Mock Midterm I: Programming Part

* Answer each questions as instructed. 

* When you finish the exam:
    * Step 1: Submit the print-out to intructor/TA.
    * Step 2: Save the file.
    * Step 3: Download the file `.ipynb`:
    * Step 4: Submission to Autograder: https://iu.instructure.com/courses/2284109/assignments/17395296

## p1.(1) [1 pts]

```python
p1_1 = None
```

```python
# check the answer
print(p1_1)
# expected to get [6, ...., 9998, 10000]
```

## p1.(2)  [1 pts]

```python
def p1_2(a):
    # finish the function:
    pass
```

```python
print(p1_2("HEllo"))      # expected output: 0
```

```python
print(p1_2("happy"))      # expected output: 1
```

```python
print(p1_2("BananA"))     # expected output: 3
```

## p1.(3).  [1 pts]

```python
def p1_3(a, b):
    # finish the function:
    pass
```

```python
p1_3(22, 25) # expected to get 45
```

```python
p1_3(1, 2) # expected to get 0
```

```python
p1_3(5, 5) # expected to get 0
```

```python
p1_3(5, 10) # expected to get 10
```

## p1.(4)  [1 pts]

```python
def p1_4(a, b):
    # finish the function:
    pass
```

```python
p1_4("abc", "abc") # expected to get 1
```

```python
p1_4("ababc", "ab") # expected to get 2
```

```python
p1_4("ababaaa", "a") # expected to get 5
```

```python
p1_4("ab", "abab") # expected to get 0
```

```python
p1_4("abc", "") # expected to get 0
```

```python

```

## p2.(1)  [1 pts]

```python
def taxi_fare(dist):
    # finish the function:
    pass
```

```python
taxi_fare(3) # expected to get 11.5
```

```python
taxi_fare(25457) # expected to get 63646.5
```

```python

```

## P2.(2)  [1pts]

```python
def total_taxi_fares(data):
    # finish the function:
    pass
```

```python
my_data_1 = [[3, 100], [103, 200], [198, 300]]
total_taxi_fares(my_data_1) # expected to return 752.0
```

```python
my_data_2 = [[0, 50]]
total_taxi_fares(my_data_2) # expected to return 129.0
```

```python
my_data_4 = [[0, 1], [1, 2], [2, 3]]
total_taxi_fares(my_data_4) # expected to return 19.5
```

Tests for p1.(1) (points: 1)
```
>>> assert p1_1 == [i for i in range(6, 10001) if i % 2 == 0]
```


Tests for p1.(2) (points: 1)
```
>>> assert p1_2('HEllo') == 0, 'Test case 01 failed'
```
```
>>> assert p1_2('happy') == 1, 'Test case 02 failed'
```
```
>>> assert p1_2('BananA') == 3, 'Test case 03 failed'
```
```
>>> test_cases = [('HEllo', 0), ('AAAA', 4), ('aaaa', 4), ('BananA', 3), ('', 0), ('123!@#', 0), ('A a A a', 4), ('BbCcDdEe', 0)]
>>> for (i, (input_str, expected_output)) in enumerate(test_cases):
...     result = p1_2(input_str)
...     assert result == expected_output, f"Test case {i + 1} failed: input='{input_str}', expected={expected_output}, got={result}"
```


Tests for p1.(3) (points: 1)
```
>>> assert p1_3(22, 25) == 47, 'Test case 1 failed'
>>> assert p1_3(1, 2) == 0, 'Test case 2 failed'
>>> assert p1_3(5, 5) == 0, 'Test case 3 failed'
>>> assert p1_3(5, 10) == 10, 'Test case 4 failed'
```
```
>>> test_cases = [(22, 25, 47), (1, 2, 0), (5, 10, 10), (5, 5, 0), (6, 3, 6), (-10, 7, 7), (0, 0, 0), (1000000, 6, 1000006)]
>>> for (i, (a, b, expected_output)) in enumerate(test_cases):
...     result = p1_3(a, b)
...     assert result == expected_output, f'Test case {i + 1} failed: a={a}, b={b}, expected={expected_output}, got={result}'
```


Tests for p1.(4) (points: 1)
```
>>> assert p1_4('abc', 'abc') == 1, 'Test case 1 failed'
>>> assert p1_4('ababc', 'ab') == 2, 'Test case 2 failed'
>>> assert p1_4('ababaaa', 'a') == 5, 'Test case 3 failed'
>>> assert p1_4('ab', 'abab') == 0, 'Test case 4 failed'
>>> assert p1_4('abc', '') == 0, 'Test case 5 failed'
```
```
>>> additional_test_cases = [('aaaa', 'aa', 2), ('abcabc', 'abc', 2), ('a$b$c$', '$', 3), ('hello', 'l', 2), ('short', 'longer', 0), ('identical', 'identical', 1), ('aaa', 'aa', 1), ('hello world', ' ', 1), ('Hello', 'hello', 0), ('abcdef', 'xyz', 0)]
>>> for (i, (a, b, expected_output)) in enumerate(additional_test_cases):
...     result = p1_4(a, b)
...     assert result == expected_output, f"Additional test case {i + 1} failed: a='{a}', b='{b}', expected={expected_output}, got={result}"
```


Tests for p2.(1) (points: 1)
```
>>> assert taxi_fare(3) == 11.5, 'Test case 1 failed'
>>> assert taxi_fare(25457) == 63646.5, 'Test case 2 failed'
```
```
>>> assert taxi_fare(3) == 11.5, 'Test case 1 failed'
>>> assert taxi_fare(25457) == 63646.5, 'Test case 2 failed'
>>> assert taxi_fare(0) == 4.0, 'Test case 3 failed'
>>> assert taxi_fare(1) == 6.5, 'Test case 4 failed'
>>> assert taxi_fare(10) == 29.0, 'Test case 5 failed'
>>> assert taxi_fare(0.5) == 5.25, 'Test case 6 failed'
```


Tests for p2.(2) (points: 1)
```
>>> test_cases = [([[10, 20]], 29.0), ([[0, 0]], 4.0), ([[3, 100], [103, 200], [198, 300]], 752.0), ([], 0.0), ([[0.5, 10.5]], 29.0), ([[0, 100000]], 250004.0), ([[5, 5]], 4.0), ([[0.0001, 0.0002]], 4.00025), ([[10, 10], [20, 20], [30, 30]], 12.0)]
>>> for (i, (data, expected_output)) in enumerate(test_cases):
...     if data and any((end < start for (start, end) in data)):
...         try:
...             total_taxi_fares(data)
...             print(f'Test case {i + 1} failed: Expected invalid input handling for negative distance.')
...         except ValueError:
...             print(f'Test case {i + 1} passed: Negative distance handled correctly.')
...     else:
...         result = total_taxi_fares(data)
...         assert result == expected_output, f'Test case {i + 1} failed: data={data}, expected={expected_output}, got={result}'
...         print(f'Test case {i + 1} passed: data={data}, expected={expected_output}, got={result}')
Test case 1 passed: data=[[10, 20]], expected=29.0, got=29.0
Test case 2 passed: data=[[0, 0]], expected=4.0, got=4.0
Test case 3 passed: data=[[3, 100], [103, 200], [198, 300]], expected=752.0, got=752.0
Test case 4 passed: data=[], expected=0.0, got=0
Test case 5 passed: data=[[0.5, 10.5]], expected=29.0, got=29.0
Test case 6 passed: data=[[0, 100000]], expected=250004.0, got=250004.0
Test case 7 passed: data=[[5, 5]], expected=4.0, got=4.0
Test case 8 passed: data=[[0.0001, 0.0002]], expected=4.00025, got=4.00025
Test case 9 passed: data=[[10, 10], [20, 20], [30, 30]], expected=12.0, got=12.0
```
