# CSCI-A201/597 Introduction to Programming I: 
# Midterm Exam: Programming Part

* Answer each questions as instructed. 

* When you finish the exam:
    * Step 1: Submit the print-out to intructor/TA.
    * Step 2: Save the file.
    * Step 3: Download the file `.ipynb`:
        * Open a new window in your browser with this link: https://ctshare.github.io/CSCIA201202/tree/.
        * You will see the window as the image below. Select the file `**.ipynb` and click botton `Download`. The file will be saved to your device.
    * Step 4: Submission to Autograder: https://iu.instructure.com/courses/2284109/assignments/17304220. 
    
<img src="plugin/ex01.png" width="350"/>

```python
# Please run this cell before programming!
# You will see a button "Click This Button Before Typing Aything"==> click the button!
from plugin import local_update 
```

## p1.(1).  [10 pts]

```python
# finish the program: 
p1_1 = None
```

```python
print(p1_1) # expected to print [5, 6, 7, ...., 9999, 10000]
```

## p1.(2).  [10 pts]

```python
# finish the function:

def p1_2(a):
    pass
```

```python
p1_2("HEllo") # expected to return 1
```

```python
p1_2("GOOdbye") #expected to return 2
```

```python
p1_2("happy") #expected to return 0
```

```python
p1_2("123") #expected to return 0
```

## p1.(3).  [10 pts]

```python
# finish the function:

def p1_3(a, b, c):
    pass
```

```python
p1_3(22, 25, 28) # expected to return 0
```

```python
p1_3(20, 29, 25) # expected to return 0
```

```python
p1_3(20, 29, 25) # expected to return 0
```

```python
p1_3(10, 15, 30) # expected to return 55
```

```python
p1_3(20, 12, 18) # expected to return  30
```

## p1.(4)  [10 pts]

```python
# finish the function:

def p1_4(a, b):
    pass
```

```python
p1_4("abc", "abc") # expected to return 2
```

```python
p1_4("abc", "xyz") # expected to return 0
```

```python
p1_4("docatzz", "dobatz") # expected to return 3
```

```python
p1_4("aaaaaa", "aaaa") # expected to return 3
```

```python
p1_4("", "xyz") # expected to return 0
```

```python

```

### P2 [20 pts]

```python
# P2 (1)


# write your code here

```

```python
####sample test 2
initial, final = 100.0, 100.0
inflation_rate(initial, final) # expected to return 0.0
```

```python
####sample test 2
initial, final = 99.99, 100.01
inflation_rate(initial, final) # expected to return 0.02
```

```python
####sample test 3
initial, final = 100.0, 80.0
inflation_rate(initial, final) # expected to return -20.0
```

```python
# P2 (2)


# write your code here

```

```python
####sample test 1
my_data = [[23.5, 24.00], [24.0, 24.25], [24.25, 24.38]]
average_inflation_rate(my_data) # expected to return 1.24
```

```python
####sample test 2
data1 = [[23.5, 24.0], [24.0, 24.25], [24.25, 24.38], [24.38, 24.75]]
average_inflation_rate(data1) # expected to return 1.31
```

```python
####sample test 3
data2 = [[25.0, 25.0], [24.0, 24.0], [30.0, 30.0], [27.0, 27.0], [50.0, 50.0]]
average_inflation_rate(data2) # expected to return 0.00
```

## p3  [10 pts]

```python
# write your code here:




```

```python
# sample test 1: 
attendance_award("PPALLP") # expected to return True 
```

```python
# sample test 2: 
attendance_award("PPALLL") # expected to return False 
```

```python
# sample test 3: 
attendance_award("PPPP") # expected to return False 
```

```python

```

## p4  [10 pts]

```python
# write your code here:




```

```python
print(factorial(0)) # Expected output: 1
print(factorial(1)) # Expected output: 1
print(factorial(3)) # Expected output: 6
print(factorial(5)) # Expected output: 120
print(factorial(20)) # Expected output: 2432902008176640000
```

```python
print(sum_multiples_35(10))  # Expected output: 23
print(sum_multiples_35(4))   # Expected output: 3
print(sum_multiples_35(3))   # Expected output: 0
print(sum_multiples_35(1))   # Expected output: 0
```

```python
print(calculator(factorial, 5))                # Expected output: 120
print(calculator(sum_multiples_35, 5))         # Expected output: 3
print(calculator(factorial, 1))                # Expected output: 1
print(calculator(sum_multiples_35, 1))         # Expected output: 0
print(calculator(lambda n: n**2, 5))           # Expected output: 25
print(calculator(lambda n: n**2, 1))           # Expected output: 1
```

## p5  [10 pts]

```python
# write your code here:




```

```python
############sample test 1: 
m1 = [[1, 2, 3],
       [4, 5, 6]]
m1t = [[1, 4], 
       [2, 5], 
       [3, 6]]
transpose(m1) == m1t # Expected output: True
```

```python
#############sample test 2: 
m2 = [[1, 2],
       [4, 5]]
m2t = [[1, 4], 
         [2, 5]]
transpose(m2) == m2t # Expected output: True
```

```python
#############sample test 3: 
m3 = [[1]]
m3t = [[1]]
transpose(m3) == m3t # Expected output: True
```

```python
#############sample test 4: 
m4 = [[1, 2]]
m4t = [[1], [2]]
transpose(m4) == m4t # Expected output: True
```

```python

```

Tests for p1.(1) (points: 10)
```
>>> def check_p1_1():
...     assert isinstance(p1_1, list), 'not a list'
...     assert len(p1_1) == 9996
...     for i in range(5, 10001):
...         assert i in p1_1, '{i} should be included'
>>> check_p1_1()
```


Tests for p1.(2) (points: 10)
```
>>> assert p1_2('AEIOU') == 5, 'Test Case 1 Failed'
```
```
>>> assert p1_2('HELLO WORLD') == 3, 'Test Case 2 Failed'
```
```
>>> assert p1_2('BCDFG') == 0, 'Test Case 3 Failed'
```
```
>>> assert p1_2('') == 0, 'Test Case 4 Failed'
```
```
>>> assert p1_2('HEllo') == 1, 'Test Case 5 Failed'
```
```
>>> assert p1_2('happy') == 0, 'Test Case 6 Failed'
```
```
>>> assert p1_2('SUNNY') == 1, 'Test Case 7 Failed'
```
```
>>> assert p1_2('123') == 0, 'Test Case 8 Failed'
```
```
>>> assert p1_2('BCDFGHJKLMNPQRSTVWXYZ' * 100) == 0, 'Test Case 9 Failed'
```


Tests for p1.(3) (points: 10)
```
>>> assert p1_3(22, 25, 28) == 0, 'Test Case 1 Failed'
```
```
>>> assert p1_3(20, 29, 25) == 0
```
```
>>> assert p1_3(10, 15, 30) == 55, 'Test Case 3 Failed'
```
```
>>> assert p1_3(20, 21, 29) == 0, 'Test Case 4 Failed'
```
```
>>> assert p1_3(20, 12, 18) == 30, 'Test Case 5 Failed'
```
```
>>> assert p1_3(20, 25, 50) == 50, 'Test Case 6 Failed'
```
```
>>> assert p1_3(5, 15, 35) == 55, 'Test Case 7 Failed'
```


Tests for p1.(4) (points: 10)
```
>>> assert p1_4('abc', 'abc') == 2, f"Failed Test Case 1: p1_4('abc', 'abc') should return 2"
```
```
>>> assert p1_4('abc', 'xyz') == 0, f"Failed Test Case 2: p1_4('abc', 'xyz') should return 0"
```
```
>>> assert p1_4('aabbcc', 'abbcdd') == 0, f"Failed Test Case 3: p1_4('aabbcc', 'abbcdd') should return 0"
```
```
>>> assert p1_4('abc', 'ab') == 1, f"Failed Test Case 4: p1_4('abc', 'ab') should return 1"
```
```
>>> assert p1_4('', 'xyz') == 0, f"Failed Test Case 5: p1_4('', 'xyz') should return 0"
```
```
>>> assert p1_4('abcdefgh', 'ijklmnop') == 0, f"Failed Test Case 6: p1_4('abcdefgh', 'ijklmnop') should return 0"
```
```
>>> assert p1_4('abcde', 'abcfg') == 2, f"Failed Test Case 7: p1_4('abcde', 'abcfg') should return 2"
```
```
>>> assert p1_4('xyzabc', 'pqrabc') == 2, f"Failed Test Case 8: p1_4('xyzabc', 'pqrabc') should return 2"
```
```
>>> assert p1_4('abcdef', '123def') == 2, f"Failed Test Case 9: p1_4('abcdef', '123def') should return 2"
```
```
>>> assert p1_4('aaaaaa', 'aaaa') == 3, f"Failed Test Case 10: p1_4('aaaaaa', 'aaaa') should return 3"
```


Tests for p2 (points: 20)
```
>>> initial = 100.0
>>> final = 100.0
>>> assert inflation_rate(initial, final) == 0.0, f'Expected 0.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 100.0
>>> final = 120.0
>>> assert inflation_rate(initial, final) == 20.0, f'Expected 20.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 100.0
>>> final = 80.0
>>> assert inflation_rate(initial, final) == -20.0, f'Expected -20.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 50.0
>>> final = 150.0
>>> assert inflation_rate(initial, final) == 200.0, f'Expected 200.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 99.99
>>> final = 100.01
>>> assert inflation_rate(initial, final) == 0.02, f'Expected 0.02, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 1000.0
>>> final = 100.0
>>> assert inflation_rate(initial, final) == -90.0, f'Expected -90.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 1000.0
>>> final = 999.99
>>> assert inflation_rate(initial, final) == 0.0
```
```
>>> my_data = [[23.5, 24.0], [24.0, 24.25], [24.25, 24.38], [24.38, 24.75]]
>>> assert average_inflation_rate(my_data) == 1.31
```
```
>>> my_data = [[23.5, 24.0], [24.0, 24.25], [24.25, 24.38]]
>>> assert average_inflation_rate(my_data) == 1.24
```
```
>>> my_data = [[25.0, 25.0], [24.0, 24.0], [30.0, 30.0], [27.0, 27.0], [50.0, 50.0]]
>>> assert average_inflation_rate(my_data) == 0.0, 'Test Case 2 Failed'
```
```
>>> my_data = [[30.0, 28.0], [25.0, 24.0], [40.0, 35.0], [50.0, 55.0], [60.0, 70.0], [80.0, 90.0]]
>>> assert average_inflation_rate(my_data) == 2.67, 'Test Case 3 Failed'
```
```
>>> my_data = [[50.0, 55.0], [55.0, 52.0], [52.0, 60.0], [60.0, 58.0], [58.0, 62.0]]
>>> assert average_inflation_rate(my_data) == 4.7, 'Test Case 5 Failed'
```
```
>>> my_data = [[1000.0, 1200.0], [1200.0, 1000.0], [1000.0, 1500.0], [1500.0, 1300.0], [1300.0, 1700.0], [1700.0, 1600.0], [1600.0, 2000.0]]
>>> assert average_inflation_rate(my_data) == 12.84, 'Test Case 6 Failed'
```
```
>>> my_data = [[1000.0, 1000.01], [1000.01, 1000.02], [1000.02, 1000.03], [1000.03, 1000.04], [1000.04, 1000.05]]
>>> assert average_inflation_rate(my_data) == 0.0, 'Edge Case Failed'
```


Tests for p3 (points: 10)
```
>>> assert attendance_award('PPALLP') == True, f"Test failed: Expected attendance_award('PPALLP') to be True, but got {attendance_award('PPALLP')}"
```
```
>>> assert attendance_award('PPALLL') == False, f"Test failed: Expected attendance_award('PPALLL') to be False, but got {attendance_award('PPALLL')}"
```
```
>>> assert attendance_award('PPPPP') == True, f"Test failed: Expected attendance_award('PPPPP') to be True, but got {attendance_award('PPPPP')}"
```
```
>>> assert attendance_award('PPLLLPP') == False, f"Test failed: Expected attendance_award('PPLLLPP') to be False, but got {attendance_award('PPLLLPP')}"
```
```
>>> assert attendance_award('PPAALLP') == False, f"Test failed: Expected attendance_award('PPAALLP') to be False, but got {attendance_award('PPAALLP')}"
```
```
>>> assert attendance_award('PPALPP') == True, f"Test failed: Expected attendance_award('PPALPP') to be True, but got {attendance_award('PPALPP')}"
```
```
>>> assert attendance_award('AAAPP') == False, f"Test failed: Expected attendance_award('AAAPP') to be False, but got {attendance_award('AAAPP')}"
```
```
>>> assert attendance_award('PLPLPLPL') == True, f"Test failed: Expected attendance_award('PLPLPLPL') to be True, but got {attendance_award('PLPLPLPL')}"
```
```
>>> assert attendance_award('PPLLLLLP') == False, f"Test failed: Expected attendance_award('PPLLLLLP') to be False, but got {attendance_award('PPLLLLLP')}"
```
```
>>> assert attendance_award('PPALLPPL') == True, f"Test failed: Expected attendance_award('PPALLPPL') to be True, but got {attendance_award('PPALLPPL')}"
```
```
>>> assert attendance_award('PPPP') == False
```
```
>>> assert attendance_award('PPPPP') == True
```
```
>>> long_string_no_absences = 'P' * 1000 + 'L' * 2 + 'P' * 1000
>>> assert attendance_award(long_string_no_absences) == True, f'Test failed: Expected attendance_award(long_string_no_absences) to be True, but got {attendance_award(long_string_no_absences)}'
```
```
>>> long_string_three_consecutive_lates = 'P' * 1000 + 'LLL' + 'P' * 1000
>>> assert attendance_award(long_string_three_consecutive_lates) == False, f'Test failed: Expected attendance_award(long_string_three_consecutive_lates) to be False, but got {attendance_award(long_string_three_consecutive_lates)}'
```


Tests for p4 (points: 30)
```
>>> assert factorial(1) == 1, f'Test failed: Expected factorial(1) to be 1, but got {factorial(1)}'
```
```
>>> assert factorial(5) == 120, f'Test failed: Expected factorial(5) to be 120, but got {factorial(5)}'
```
```
>>> assert factorial(10) == 3628800, f'Test failed: Expected factorial(10) to be 3628800, but got {factorial(10)}'
```
```
>>> assert factorial(20) == 2432902008176640000, f'Test failed: Expected factorial(20) to be 2432902008176640000, but got {factorial(20)}'
```
```
>>> assert factorial(3) == 6, f'Test failed: Expected factorial(3) to be 6, but got {factorial(3)}'
```
```
>>> assert factorial(15) == 1307674368000, f'Test failed: Expected factorial(15) to be 1307674368000, but got {factorial(15)}'
```
```
>>> assert factorial(17) == 355687428096000, f'Test failed: Expected factorial(17) to be 355687428096000, but got {factorial(17)}'
```
```
>>> assert factorial(23) == 25852016738884976640000, f'Test failed: Expected factorial(23) to be 25852016738884976640000, but got {factorial(23)}'
```
```
>>> assert factorial(30) == 265252859812191058636308480000000, f'Test failed: Expected factorial(30) to be 265252859812191058636308480000000, but got {factorial(30)}'
```
```
>>> assert sum_multiples_35(1) == 0, f'Test failed: Expected sum_multiples_35(1) to be 0, but got {sum_multiples_35(1)}'
```
```
>>> assert sum_multiples_35(3) == 0, f'Test failed: Expected sum_multiples_35(3) to be 0, but got {sum_multiples_35(3)}'
```
```
>>> assert sum_multiples_35(2) == 0, f'Test failed: Expected sum_multiples_35(2) to be 0, but got {sum_multiples_35(2)}'
```
```
>>> assert sum_multiples_35(4) == 3, f'Test failed: Expected sum_multiples_35(4) to be 3, but got {sum_multiples_35(3)}'
```
```
>>> assert sum_multiples_35(6) == 8, f'Test failed: Expected sum_multiples_35(6) to be 8, but got {sum_multiples_35(6)}'
```
```
>>> assert sum_multiples_35(10) == 23, f'Test failed: Expected sum_multiples_35(10) to be 23, but got {sum_multiples_35(10)}'
```
```
>>> assert sum_multiples_35(15) == 45, f'Test failed: Expected sum_multiples_35(15) to be 45, but got {sum_multiples_35(15)}'
```
```
>>> assert sum_multiples_35(50) == 543, f'Test failed: Expected sum_multiples_35(50) to be 543, but got {sum_multiples_35(50)}'
```
```
>>> assert sum_multiples_35(100) == 2318, f'Test failed: Expected sum_multiples_35(100) to be 2318, but got {sum_multiples_35(100)}'
```
```
>>> assert calculator(factorial, 5) == 120, f'Test failed: Expected calculator(factorial, 5) to return 120, but got {calculator(factorial, 5)}'
```
```
>>> assert calculator(sum_multiples_35, 5) == 3, f'Test failed: Expected calculator(sum_multiples_35, 5) to return 3, but got {calculator(sum_multiples_35, 5)}'
```
```
>>> assert calculator(factorial, 1) == 1, f'Test failed: Expected calculator(factorial, 1) to return 1, but got {calculator(factorial, 1)}'
```
```
>>> assert calculator(sum_multiples_35, 1) == 0, f'Test failed: Expected calculator(sum_multiples_35, 1) to return 0, but got {calculator(sum_multiples_35, 1)}'
```
```
>>> assert calculator(factorial, 10) == 3628800, f'Test failed: Expected calculator(factorial, 10) to return 3628800, but got {calculator(factorial, 10)}'
```
```
>>> assert calculator(lambda n: n + 10, 5) == 15, f'Test failed: Expected calculator(lambda n: n + 10, 5) to return 15, but got {calculator(lambda n: n + 10, 5)}'
```
```
>>> assert calculator(lambda n: n * 3, 7) == 21, f'Test failed: Expected calculator(lambda n: n * 3, 7) to return 21, but got {calculator(lambda n: n * 3, 7)}'
```
```
>>> assert calculator(lambda n: n ** 3, 4) == 64, f'Test failed: Expected calculator(lambda n: n ** 3, 4) to return 64, but got {calculator(lambda n: n ** 3, 4)}'
```
```
>>> assert calculator(lambda n: n % 2, 9) == 1, f'Test failed: Expected calculator(lambda n: n % 2, 9) to return 1, but got {calculator(lambda n: n % 2, 9)}'
```
```
>>> assert calculator(lambda n: n ** 2 + 2 * n + 6, 10) == 126
```
```
>>> assert calculator(lambda n: 2 * n ** 3 + 2 * n + 6, 10) == 2026
```
```
>>> assert calculator(lambda n: n % 2 == 0, 8) == True, f'Test failed: Expected calculator(lambda n: n % 2 == 0, 8) to return True, but got {calculator(lambda n: n % 2 == 0, 8)}'
```


Tests for p5 (points: 10)
```
>>> original_matrix_1 = [[1, 2, 3], [4, 5, 6], [7, 8, 9]]
>>> expected_output_1 = [[1, 4, 7], [2, 5, 8], [3, 6, 9]]
>>> assert transpose(original_matrix_1) == expected_output_1
```
```
>>> original_matrix_2 = [[1, 2], [3, 4], [5, 6]]
>>> expected_output_2 = [[1, 3, 5], [2, 4, 6]]
>>> assert transpose(original_matrix_2) == expected_output_2
```
```
>>> original_matrix_3 = [[1, 2, 3], [4, 5, 6]]
>>> expected_output_3 = [[1, 4], [2, 5], [3, 6]]
>>> assert transpose(original_matrix_3) == expected_output_3
```
```
>>> original_matrix_4 = [[42]]
>>> expected_output_4 = [[42]]
>>> assert transpose(original_matrix_4) == expected_output_4
```
```
>>> original_matrix_5 = [[]]
>>> expected_output_5 = [[]]
>>> assert transpose(original_matrix_5) == expected_output_5
```
```
>>> original_matrix_6 = [[1, 2, 3]]
>>> expected_output_6 = [[1], [2], [3]]
>>> assert transpose(original_matrix_6) == expected_output_6
```
```
>>> original_matrix_7 = [[1], [2], [3]]
>>> expected_output_7 = [[1, 2, 3]]
>>> assert transpose(original_matrix_7) == expected_output_7
```
```
>>> original_matrix_7 = [[1]]
>>> expected_output_7 = [[1]]
>>> assert transpose(original_matrix_7) == expected_output_7
```
