# CSCI-A201/597 Introduction to Programming I: 
# Midterm Exam: Programming Part

* Answer each questions as instructed. 
* When you finish the exam:
    * Before submit the file to Autograder: submit the print-out to intructor/TA.
    * Download the notebook file (`.ipynb`) to Autograder on Canvas. 

### P1.  [20 pts]


```python
# P1 (1)

# finish the program: 
p1_1 = [i for i in range(5, 101)] # solution
```

```python
# p1 (2)
# finish the function:

def p1_2(a, b, c):
    pass
```

```python
# p1 (3)
# finish the function:

def p1_3(a, b, c):
    pass
```

```python
# p1 (4)
# finish the function:

def p1_3(a, b):
    pass
```

```python

```

### P2  [20 pts]

```python
# P2 (1)


# write your code here

```

```python
# P2 (2)


# write your code here

```

```python

```

Tests for p1 (points: 20)
```
>>> def check_p1_1():
...     assert isinstance(p1_1, list), 'not a list'
...     assert len(p1_1) == 96
...     for i in range(5, 101):
...         assert i in p1_1, '{i} should be included'
>>> check_p1_1()
```
```
>>> def test_p1_2():
...     assert p1_2(3, 4, 7) == True, f'Failed Test Case 1: p1_2(3, 4, 7) should return True'
...     assert p1_2(7, 10, 6) == True, f'Failed Test Case 2: p1_2(7, 10, 6) should return True'
...     assert p1_2(1, 4, 6) == False, f'Failed Test Case 3: p1_2(1, 4, 6) should return False'
...     assert p1_2(1, 2, 3) == False, f'Failed Test Case 4: p1_2(1, 2, 3) should return False'
...     assert p1_2(10, 13, 9) == True, f'Failed Test Case 5: p1_2(10, 13, 9) should return True'
...     assert p1_2(5, 5, 5) == False, f'Failed Test Case 6: p1_2(5, 5, 5) should return False'
>>> test_p1_2()
```
```
>>> def test_p1_3():
...     assert p1_3(22, 25, 28) == 0, f'Failed Test Case 1: p1_3(22, 25, 28) should return 0'
...     assert p1_3(10, 15, 30) == 55, f'Failed Test Case 2: p1_3(10, 15, 30) should return 55'
...     assert p1_3(20, 12, 18) == 30, f'Failed Test Case 3: p1_3(20, 12, 18) should return 30'
...     assert p1_3(27, 30, 25) == 30, f'Failed Test Case 4: p1_3(27, 30, 25) should return 30'
...     assert p1_3(5, 10, 15) == 30, f'Failed Test Case 5: p1_3(5, 10, 15) should return 30'
...     assert p1_3(20, 29, 25) == 0, f'Failed Test Case 6: p1_3(20, 29, 25) should return 0'
...     assert p1_3(23, 17, 35) == 52, f'Failed Test Case 7: p1_3(23, 17, 35) should return 52'
>>> test_p1_3()
```
```
>>> def test_p1_4():
...     assert p1_4('abc', 'abc') == 2, f"Failed Test Case 1: p1_4('abc', 'abc') should return 2"
...     assert p1_4('abc', 'xyz') == 0, f"Failed Test Case 2: p1_4('abc', 'xyz') should return 0"
...     assert p1_4('aabbcc', 'abbcdd') == 0, f"Failed Test Case 3: p1_4('aabbcc', 'abbcdd') should return 0"
...     assert p1_4('abc', 'ab') == 1, f"Failed Test Case 4: p1_4('abc', 'ab') should return 1"
...     assert p1_4('', 'xyz') == 0, f"Failed Test Case 5: p1_4('', 'xyz') should return 0"
...     assert p1_4('abcdefgh', 'ijklmnop') == 0, f"Failed Test Case 6: p1_4('abcdefgh', 'ijklmnop') should return 0"
...     assert p1_4('abcde', 'abcfg') == 2, f"Failed Test Case 7: p1_4('abcde', 'abcfg') should return 2"
...     assert p1_4('xyzabc', 'pqrabc') == 2, f"Failed Test Case 8: p1_4('xyzabc', 'pqrabc') should return 2"
...     assert p1_4('abcdef', '123def') == 2, f"Failed Test Case 9: p1_4('abcdef', '123def') should return 2"
...     assert p1_4('aaaaaa', 'aaaa') == 3, f"Failed Test Case 10: p1_4('aaaaaa', 'aaaa') should return 3"
>>> test_p1_4()
```


Tests for p2 (points: 20)
```
>>> def test_inflation_rate():
...     initial = 100.0
...     final = 100.0
...     assert inflation_rate(initial, final) == 0.0, f'Expected 0.0, but got {inflation_rate(initial, final)}'
...     initial = 100.0
...     final = 120.0
...     assert inflation_rate(initial, final) == 20.0, f'Expected 20.0, but got {inflation_rate(initial, final)}'
...     initial = 100.0
...     final = 80.0
...     assert inflation_rate(initial, final) == -20.0, f'Expected -20.0, but got {inflation_rate(initial, final)}'
...     initial = 50.0
...     final = 150.0
...     assert inflation_rate(initial, final) == 200.0, f'Expected 200.0, but got {inflation_rate(initial, final)}'
...     initial = 99.99
...     final = 100.01
...     assert inflation_rate(initial, final) == 0.02, f'Expected 0.02, but got {inflation_rate(initial, final)}'
...     initial = 1000.0
...     final = 100.0
...     assert inflation_rate(initial, final) == -90.0, f'Expected -90.0, but got {inflation_rate(initial, final)}'
>>> test_inflation_rate()
```
```
>>> def test_average_inflation_rate():
...     expected = 1.24
...     result = average_inflation_rate()
...     assert result == expected, f'Test Case 1 Failed: expected {expected}, got {result}'
>>> test_average_inflation_rate()
```
