# CSCI-A202/598 Introduction to Programming II: 
# Midterm II: Programming Part

* Answer each questions as instructed. 

* When you finish the exam:
    * Step 1: Submit the print-out to intructor/TA.
    * Step 2: Save the file.
    * Step 3: Download the file `.ipynb`:
    * Step 4: Submission to Autograder

## p1.(1) [20 pts]

Create a `Furniture` class that has the following properties:
* `name` : a string representing the name of the furniture.
* `area` : an integer representing the area occupied by the furniture in sqft. An `AssertionError` is raised if the area is not an integer.
* a special method that when a `Furniture` object is printed, a string will be printout such as: `"Bed:200 sqft".
* a method `get_area` that return a string to show the value of `area`.

```python
class Furniture:
    '''
    A class to represent furniture items.
    >>> chair = Furniture("Chair", 50)
    >>> print(chair)
    Chair:50 sqft
    >>> chair.get_area()
    '50 sqft'
    >>> table = Furniture("Table", "100")
    Traceback (most recent call last):
    ...
    AssertionError: Area must be an integer
    '''
    "*** YOUR CODE HERE ***"
    
```

```python
# Create furniture objects
bed = Furniture("Bed", 200)
print(bed)  # Output: "Bed:200 sqft"
print(bed.get_area())  # Output: "200 sqft"
```

```python
# This will raise an AssertionError
sofa = Furniture("Sofa", "150")  # Raises: AssertionError: Area must be an integer
```

## P1.(2) [20 pts]

**Hourse**

Create a class named `House` which maintains a list of Furniture objects
with their respective areas. The class has the followed attributes and methods:
* Attributes:
    * `furniture_list` : a list that contains all the furniture objects in the house
    * `area` : a integer that represents the total area of the house, with a default value 0.

* Methods:
    * `add_furniture` : a method that adds a `Furniture` object to the `furniture_list`. If the total area by the furniture exceeds 50% of the area of the house, an exception is raised.
    * A special method `__str__` that returns a string representation of the `House` object, showing the list of furniture as in examples followed.

```python
class House:
    '''
    >>> house = House(1000)
    >>> bed = Furniture("Bed", 200)
    >>> chair = Furniture("Chair", 50)
    >>> house.add_furniture(bed)
    >>> house.add_furniture(chair)
    >>> print(house)
    House(1000 sqft) with furniture: Bed:200 sqft, Chair:50 sqft
    >>> sofa = Furniture("Sofa", 300)
    >>> house.add_furniture(sofa)  # Total would be 550/1000 (55%)
    Traceback (most recent call last):
    ...
    Exception: Furniture exceeds 50% of house area
    '''
    "*** YOUR CODE HERE ***"
```

```python
# Create a house with 800 sqft
my_house = House(800)

# Add some furniture
my_house.add_furniture(Furniture("Bed", 200))
my_house.add_furniture(Furniture("Desk", 100))

print(my_house)
# Output: House(800 sqft) with furniture: Bed:200 sqft, Desk:100 sqft
```

```python
# This will raise an exception
my_house.add_furniture(Furniture("Sofa", 300))  
# Raises: Exception: Furniture exceeds 50% of house area
```

## p2.  [20 pts]

**Linked List**

You're provided with the `Link` class as followed:

```python
class Link(object):
    empty = ()

    def __init__(self, first, rest=empty):
        self.first = first
        self.rest = rest

    def __len__(self):
        return 1 + len(self.rest)

    def __getitem__(self, i):
        if i == 0:
            return self.first
        return self.rest[i - 1]

    def __repr__(self):
        if self.rest is empty:
            return 'Link({})'.format(repr(self.first))
        return 'Link({}, {})'.format(repr(self.first), repr(self.rest))
```

Implement a function `count`, which takes a `Link` and another value, and counts the number of times that value is found in the `Link`.

```python
def count(r, value):
    '''Counts the number of times VALUE shows up in R.

    >>> r = Link(3, Link(3, Link(2, Link(3))))
    >>> count(r, 3)
    3
    >>> count(r, 2)
    1
    '''
    "*** YOUR CODE HERE ***"
```

```python
r = Link(3, Link(3, Link(2, Link(3))))
count(r, 3) # return 3
```

```python
count(r, 2) # return 2
```

```python
link_no_match = Link(1, Link(2, Link(3, Link(4))))
count(link_no_match, 0) # return 0
```

```python
link_mixed = Link("a", Link(1, Link("a")))
count(link_mixed, "a") # return 2
```

## p3.(1)  [20 pts]

You're provided with a `Time` class as followed:

```python
class Time:
    
    time_zone = None
    
    def __init__(self, h=0, m=0, s=0):
        for i in [h, m, s]:
            if not isinstance(i, int):
                raise TypeError("Integers must be used to represent hours, minutes, and seconds in time.")
        if not -1 < h < 24:
            raise ValueError("Invalid time range!")
        for i in [m, s]:
            if not -1 < i < 60:
                raise ValueError("Invalid time range!")   
        self.hours = h
        self.minutes = m
        self.seconds = s
    
    def __str__(self):
        return "{}:{}:{}".format(self.hours, self.minutes, self.seconds)
        
    def __repr__(self):
        return "Time({}, {}, {})".format(self.hours, self.minutes, self.seconds)
    
    def get_hours(self):
        return str(self.hours)
    
    def get_minutes(self):
        return str(self.minutes)
    
    def get_seconds(self):
        return str(self.seconds)
```

**Implementing the Bloomington_Time Class**
You will create a `Bloomington_Time` class that inherits from the provided `Time` class. 

Your task is to extend this class to represent time specifically in Bloomington, Indiana, by setting a fixed time zone, and to add functionality for adding two times (+) and comparing them (<=). This will make `Bloomington_Time` suitable for time arithmetic and scheduling tasks in a specific time zone.

* Time Zone: Override the `time_zone` class attribute to `"EST"` (Eastern Standard Time, as Bloomington, Indiana, typically uses EST, UTC-5).
* Initialization: Use the parent class’s `__init__` method to set hours, minutes, and seconds, ensuring inherited validation works.
* String Representation: Modify `__str__` to include the time zone (e.g., "EST, 12:30:45").
* Addition (+): Implement `__add__` to add two `Bloomington_Time` objects, returning a new `Bloomington_Time` object with the summed time, handling carry-over and wrapping at 24 hours.
* Comparison (<=): Implement `__le__` to compare two `Bloomington_Time` objects chronologically, returning `True` if the current time is less than or equal to the other.

```python
class Bloomington_Time(Time):
    "*** YOUR CODE HERE ***"
```

```python
t1 = Bloomington_Time(12, 30, 45)
print(t1)  # Expected: "EST, 12:30:45"

t2 = Bloomington_Time()  # Default time
print(t2)  # Expected: "EST, 0:0:0"
```

```python
t3 = Bloomington_Time(21, 30, 45)
t4 = Bloomington_Time(1, 40, 20)
t5 = t3 + t4  # 21:30:45 + 1:40:20 = 23:11:05
print(t5)  # Expected: "EST, 23:11:5"

t6 = Bloomington_Time(23, 50, 40)
t7 = Bloomington_Time(0, 20, 30)
t8 = t6 + t7  # 23:50:40 + 0:20:30 = 0:11:10
print(t8)  # Expected: "EST, 0:11:10"
```

```python
t9 = Bloomington_Time(12, 0, 0)
t10 = Bloomington_Time(13, 0, 0)
print(t9 <= t10)  # Expected: True (12:00:00 ≤ 13:00:00)
print(t10 <= t9)  # Expected: False (13:00:00 ≤ 12:00:00)

t11 = Bloomington_Time(14, 30, 0)
t12 = Bloomington_Time(14, 30, 0)
print(t11 <= t12)  # Expected: True (14:30:00 ≤ 14:30:00)
```

```python

```

### P3.(2) [20 pts]

In this task, you will create a `Clock` class that simulates a digital clock in Bloomington, Indiana, using the provided `Bloomington_Time` class to manage time values. The `Clock` class should function as an infinity **iterator**, yielding the current time at each step and advancing it by a specified duration. Each call to `next()` on the iterator should print the current time and increment it, allowing users to track time progression step-by-step.

```python
class Clock:
    def __init__(self, start_time, duration = Bloomington_Time(0, 10, 0)):
        '''Initialize the Clock with a start time and optional duration.'''
        "*** YOUR CODE HERE ***"
        
    def __iter__(self):
        '''Return the iterator object.'''
        "*** YOUR CODE HERE ***"
        
    def __next__(self):
        '''Increment the current time and advance it by the duration.'''
        "*** YOUR CODE HERE ***"
   
```

```python
start = Bloomington_Time(17, 30, 0)
clock = Clock(start)  # Default 10-minute duration
clock_iter = iter(clock)

time1 = next(clock_iter)  
print(time1)              # Prints: "EST, 17:30:0", advances to 17:40:0
time2 = next(clock_iter)  
print(time2)              # Prints: "EST, 17:40:0", advances to 17:50:0
time3 = next(clock_iter)  
print(time3)              # Prints: "EST, 17:50:0", advances to 18:00:0
```

```python
clock = Clock(Bloomington_Time(23, 50, 0), Bloomington_Time(0, 15, 0))  # 15-minute duration
for i, time in enumerate(clock):
    print(time)
    if i >= 3:  # Limit to 3+1=4 steps
        break
# Expected output:
# EST, 23:50:0
# EST, 0:5:0  (wraps at 24 hours)
# EST, 0:20:0
# EST, 0:35:0
```

```python

```

Tests for p1.(1) (points: 20)
```
>>> chair = Furniture('Chair', 50)
>>> assert chair.name == 'Chair', "Name should be 'Chair'"
>>> assert chair.area == 50, 'Area should be 50'
```
```
>>> table = Furniture('Table', 100)
>>> assert str(table) == 'Table:100 sqft', "String representation should be 'Table:100 sqft'"
```
```
>>> sofa = Furniture('Sofa', 75)
>>> assert sofa.get_area() == '75 sqft', "get_area should return '75 sqft'"
```
```
>>> try:
...     Furniture('Desk', '100')
...     assert False, 'Expected AssertionError for non-integer area'
... except AssertionError as e:
...     assert str(e) == 'Area must be an integer', "Error message should be 'Area must be an integer'"
```
```
>>> bed = Furniture('Bed', -10)
>>> assert bed.area == -10, 'Negative area should be allowed and set to -10'
>>> assert str(bed) == 'Bed:-10 sqft', 'String representation should handle negative area'
```
```
>>> lamp = Furniture('Lamp', 0)
>>> assert lamp.get_area() == '0 sqft', "get_area should return '0 sqft' for zero area"
```
```
>>> wardrobe = Furniture('Wardrobe with Drawers', 120)
>>> assert str(wardrobe) == 'Wardrobe with Drawers:120 sqft', 'String representation should handle long names'
```


Tests for p1.(2) (points: 20)
```
>>> house1 = House(1000)
>>> assert house1.area == 1000, 'House area should be 1000'
>>> assert house1.furniture_list == [], 'Furniture list should be empty initially'
```
```
>>> house2 = House(1000)
>>> bed = Furniture('Bed', 200)
>>> chair = Furniture('Chair', 50)
>>> house2.add_furniture(bed)
>>> house2.add_furniture(chair)
>>> assert len(house2.furniture_list) == 2, 'Should have 2 furniture items'
>>> assert house2.furniture_list[0] == bed, 'First item should be bed'
>>> assert house2.furniture_list[1] == chair, 'Second item should be chair'
```
```
>>> house3 = House(1000)
>>> house3.add_furniture(Furniture('Bed', 200))
>>> house3.add_furniture(Furniture('Chair', 50))
>>> assert str(house3) == 'House(1000 sqft) with furniture: Bed:200 sqft, Chair:50 sqft', 'String representation should match expected format'
```
```
>>> house4 = House(1000)
>>> house4.add_furniture(Furniture('Bed', 200))
>>> house4.add_furniture(Furniture('Chair', 50))
>>> try:
...     house4.add_furniture(Furniture('Sofa', 300))
...     assert False, 'Expected Exception for exceeding 50% area'
... except Exception as e:
...     assert str(e) == 'Furniture exceeds 50% of house area', 'Error message should match'
```
```
>>> house5 = House(500)
>>> assert str(house5) == 'House(500 sqft) with furniture: ', 'Empty house should have no furniture in string'
```
```
>>> house6 = House(600)
>>> house6.add_furniture(Furniture('Table', 300))
>>> assert len(house6.furniture_list) == 1, 'Should allow exactly 50% area'
>>> assert str(house6) == 'House(600 sqft) with furniture: Table:300 sqft', 'String should reflect one item at 50%'
```
```
>>> house7 = House(1000)
>>> house7.add_furniture(Furniture('Desk', 200))
>>> house7.add_furniture(Furniture('Lamp', 100))
>>> house7.add_furniture(Furniture('Rug', 200))
>>> assert len(house7.furniture_list) == 3, 'Should allow 3 items up to 50%'
>>> assert sum((f.area for f in house7.furniture_list)) == 500, 'Total area should be 500'
```


Tests for p2 (points: 20)
```
>>> assert count(Link.empty, 5) == 0, 'Count should be 0 for empty list'
```
```
>>> single_link = Link(42)
>>> assert count(single_link, 42) == 1, 'Count should be 1 for single matching item'
```
```
>>> single_link_no_match = Link(10)
>>> assert count(single_link_no_match, 7) == 0, 'Count should be 0 for single non-matching item'
```
```
>>> link_one_match = Link(1, Link(2, Link(3)))
>>> assert count(link_one_match, 2) == 1, 'Count should be 1 for one match in multiple items'
```
```
>>> link_multi_match = Link(5, Link(5, Link(3, Link(5))))
>>> assert count(link_multi_match, 5) == 3, 'Count should be 3 for three matches'
```
```
>>> link_no_match = Link(1, Link(2, Link(3, Link(4))))
>>> assert count(link_no_match, 0) == 0, 'Count should be 0 when no items match'
```
```
>>> link_mixed = Link('a', Link(1, Link('a')))
>>> assert count(link_mixed, 'a') == 2, 'Count should be 2 for matching strings in mixed types'
```


Tests for p3.(1) (points: 20)
```
>>> t1 = Bloomington_Time()
>>> assert t1.hours == 0, 'Default hours should be 0'
>>> assert t1.minutes == 0, 'Default minutes should be 0'
>>> assert t1.seconds == 0, 'Default seconds should be 0'
>>> assert t1.time_zone == 'EST', "Time zone should be 'EST'"
```
```
>>> t2 = Bloomington_Time(12, 30, 45)
>>> assert t2.hours == 12, 'Hours should be 12'
>>> assert t2.minutes == 30, 'Minutes should be 30'
>>> assert t2.seconds == 45, 'Seconds should be 45'
>>> assert str(t2) == 'EST, 12:30:45', "String representation should be 'EST, 12:30:45'"
```
```
>>> t3 = Bloomington_Time(23, 59, 59)
>>> assert str(t3) == 'EST, 23:59:59', 'String representation should handle max time'
```
```
>>> t4 = Bloomington_Time(10, 20, 30)
>>> t5 = Bloomington_Time(0, 15, 45)
>>> t6 = t4 + t5
>>> assert str(t6) == 'EST, 10:36:15', "10:20:30 + 0:15:45 should be 'EST, 10:36:15'"
>>> assert t6.hours == 10, 'Hours should be 10'
>>> assert t6.minutes == 36, 'Minutes should be 36'
>>> assert t6.seconds == 15, 'Seconds should be 15'
```
```
>>> t7 = Bloomington_Time(23, 50, 40)
>>> t8 = Bloomington_Time(0, 20, 30)
>>> t9 = t7 + t8
>>> assert str(t9) == 'EST, 0:11:10', "23:50:40 + 0:20:30 should wrap to 'EST, 0:11:10'"
```
```
>>> t10 = Bloomington_Time(12, 0, 0)
>>> t11 = Bloomington_Time(13, 0, 0)
>>> assert (t10 <= t11) == True, '12:00:00 should be <= 13:00:00'
>>> assert (t11 <= t10) == False, '13:00:00 should not be <= 12:00:00'
```
```
>>> t12 = Bloomington_Time(14, 30, 0)
>>> t13 = Bloomington_Time(14, 30, 0)
>>> assert (t12 <= t13) == True, '14:30:00 should be <= 14:30:00'
>>> assert (t13 <= t12) == True, '14:30:00 should be <= 14:30:00'
```
```
>>> try:
...     Bloomington_Time(12.5, 30, 45)
...     assert False, 'Expected TypeError for non-integer input'
... except TypeError as e:
...     assert str(e) == 'Integers must be used to represent hours, minutes, and seconds in time.', 'Error message should match'
```
```
>>> try:
...     Bloomington_Time(24, 0, 0)
...     assert False, 'Expected ValueError for hours out of range'
... except ValueError as e:
...     assert str(e) == 'Invalid time range!', 'Error message should match'
```
```
>>> t14 = Bloomington_Time(9, 0, 0)
>>> try:
...     t14 + 'not_a_time'
...     assert False, 'Expected TypeError for adding non-Bloomington_Time'
... except TypeError as e:
...     assert str(e) == 'Can only add another Bloomington_Time object', 'Error message should match'
```


Tests for p3.(2) (points: 20)
```
>>> start1 = Bloomington_Time(17, 30, 0)
>>> clock1 = Clock(start1)
>>> assert clock1.current_time == start1, 'Initial time should match start_time'
>>> assert str(clock1.duration) == 'EST, 0:10:0', 'Default duration should be 10 minutes'
```
```
>>> start1 = Bloomington_Time(17, 30, 0)
>>> clock1 = Clock(start1)
>>> assert isinstance(clock1.current_time, Bloomington_Time) and str(clock1.current_time) == 'EST, 17:30:0', "Initial time should be 'EST, 17:30:0'"
>>> assert isinstance(clock1.duration, Bloomington_Time) and str(clock1.duration) == 'EST, 0:10:0', "Default duration should be 'EST, 0:10:0'"
```
```
>>> clock2 = Clock(Bloomington_Time(9, 0, 0))
>>> clock_iter2 = iter(clock2)
>>> returned_time = next(clock_iter2)
>>> assert isinstance(returned_time, Bloomington_Time) and str(returned_time) == 'EST, 9:0:0', "First returned time should be 'EST, 9:0:0'"
>>> assert isinstance(clock2.current_time, Bloomington_Time) and str(clock2.current_time) == 'EST, 9:10:0', "Time should advance to 'EST, 9:10:0'"
```
```
>>> clock3 = Clock(Bloomington_Time(12, 45, 0))
>>> clock_iter3 = iter(clock3)
>>> time1 = next(clock_iter3)
>>> time2 = next(clock_iter3)
>>> time3 = next(clock_iter3)
>>> assert isinstance(time1, Bloomington_Time) and str(time1) == 'EST, 12:45:0', "First returned time should be 'EST, 12:45:0'"
>>> assert isinstance(time2, Bloomington_Time) and str(time2) == 'EST, 12:55:0', "Second returned time should be 'EST, 12:55:0'"
>>> assert isinstance(time3, Bloomington_Time) and str(time3) == 'EST, 13:5:0', "Third returned time should be 'EST, 13:5:0'"
>>> assert isinstance(clock3.current_time, Bloomington_Time) and str(clock3.current_time) == 'EST, 13:15:0', "Final time should be 'EST, 13:15:0'"
```
```
>>> clock4 = Clock(Bloomington_Time(14, 0, 0), Bloomington_Time(0, 15, 0))
>>> clock_iter4 = iter(clock4)
>>> time1 = next(clock_iter4)
>>> time2 = next(clock_iter4)
>>> assert isinstance(time1, Bloomington_Time) and str(time1) == 'EST, 14:0:0', "First returned time should be 'EST, 14:0:0'"
>>> assert isinstance(time2, Bloomington_Time) and str(time2) == 'EST, 14:15:0', "Second returned time should be 'EST, 14:15:0'"
>>> assert isinstance(clock4.current_time, Bloomington_Time) and str(clock4.current_time) == 'EST, 14:30:0', "Time should advance to 'EST, 14:30:0'"
```
```
>>> clock5 = Clock(Bloomington_Time(23, 50, 0))
>>> clock_iter5 = iter(clock5)
>>> time1 = next(clock_iter5)
>>> time2 = next(clock_iter5)
>>> assert isinstance(time1, Bloomington_Time) and str(time1) == 'EST, 23:50:0', "First returned time should be 'EST, 23:50:0'"
>>> assert isinstance(time2, Bloomington_Time) and str(time2) == 'EST, 0:0:0', "Second returned time should be 'EST, 0:0:0'"
>>> assert isinstance(clock5.current_time, Bloomington_Time) and str(clock5.current_time) == 'EST, 0:10:0', "Time should advance to 'EST, 0:10:0'"
```
```
>>> clock6 = Clock(Bloomington_Time(8, 0, 0))
>>> assert iter(clock6) is clock6, '__iter__ should return self'
```
