# CSCI-A202/598 Introduction to Programming II: 
# Midterm I: Programming Part

* Answer each questions as instructed. 

* When you finish the exam:
    * Step 1: Submit the print-out to intructor/TA.
    * Step 2: Save the file.
    * Step 3: Download the file `.ipynb`:
    * Step 4: Submission to Autograder: https://iu.instructure.com/courses/2249164/assignments/16824705. 

## p1 [10 pts]

**You Again**

Implement `again`, which takes a function `f` as an argument. The `again` function returns the smallest non-negative integer `n` for which `f(n)` is equal to `f(m)` for some non-negative `m` that is less than `n`. Assume that `f` takes non-negative integers and returns the same value for at least two different non-negative arguments.

Constraints:
* Lines numbered 2, 4, and 5 must begin with either `while` or `if`.
* Lines numbered 6 and 7 must contain either `return` or `=`.

```python
def parabola(x):
    """A parabola function (for testing the again function)."""
    return (x - 3) * (x - 6)

def vee(x):
    """A V-shaped function (for testing the again function)."""
    return abs(x - 2)

def again(f):
    """
    Return the smallest non-negative integer n such that f(n) == f(m) for some m < n.
    """
    n = 1
    # finish the program by replace the underline with your code:
    _______________
        m = 0
        ___________________
        ___________________
        ___________________
        ___________________
        n += 1
```

```python
again(parabola)  # parabola(4) == parabola(5)
# expected to return 5
```

```python
again(vee)  # vee(1) == vee(3)
# expected to return 3
```

## P2 [40 pts]
**UP AND DOWN**

Definition. Two adjacent digits in a non-negative integer are an **increase** if the left digit is smaller than the right digit, and a **decrease** if the left digit is larger than the right digit.

For example, 61127 has 2 increases (1 → 2 and 2 → 7) and 1 decrease (6 → 1).

You may use the sign function defined below in all parts of this question.

```python
def sign(x):
    """Returns the sign of x: 1 for positive, -1 for negative, and 0 for zero."""
    if x > 0:
        return 1
    elif x < 0:
        return -1
    else:
        return 0
```

### p2.(1)  [10 pts]

Implement `ramp`, which takes a non-negative integer `n` and returns whether it has more increases than decreases when reading its digits from left to right (see the definition above).

```python
def ramp(n):
    """
    Return whether non-negative integer N has more increases than decreases.
    """
    # finish the function by replace the underlines with your code:
    n, last, tally = ______, _______, ________
    while n:
        n, last, tally = _____, ______, ________
    return ________
```

```python
ramp(123)   # 2 increases (1 -> 2, 2 -> 3) and 0 decreases
# expected to return True
```

```python
ramp(1315)  # 2 increases (1 -> 3, 1 -> 5) and 1 decrease (3 -> 1)
# expected to return True
```

```python
ramp(176) # 1 increase (1 -> 7) and 1 decrease (7 -> 6)
# expected to return False
```

```python
ramp(5) # 0 increase and 0 decrease
# expected to return False
```

### p2.(2).  [10 pts]

Implement `over_under`, which takes a number `y` and returns a function that takes a number `x`. This function returns 1 if `x` is greater than `y`, 0 if `x` equals `y`, and -1 if `x` is less than `y`.

* You may not use `if`, `and`, or `or`.
* Hint: `lambda`

```python
def over_under(y):
    """
    Return a function that takes X and returns:
        -1 if X is less than Y
         0 if X is equal to Y
         1 if X is greater than Y
    """
    # finish the function by replace the underlines with your code:
    return _________________
```

```python
over_under(5)(3)  # 3 < 5
# expected to return -1
```

```python
over_under(5)(5)  # 5 == 5
# expected to return 0
```

```python
over_under(5)(7)  # 7 > 5
# expected to return 1
```

### p2.(3)  [10 pts]

The `process` function below uses `tally` and result functions to analyze all adjacent pairs of digits in a non-negative integer `n`. A `tally` function is called on each pair of adjacent digits.

```python
def process(n, tally, result):
    """Process all pairs of adjacent digits in N using functions TALLY and RESULT."""
    while n >= 10:
        tally, result = tally(n % 100 // 10, n % 10)
        n = n // 10
    return result()
```

Implement `ups`, which returns two functions that can be passed as `tally` and result arguments to process, so that process computes whether a non-negative integer `n` has exactly `k` increases.

* Hint: You can use `sign` from the previous questions and the built-in `max` and `min` functions.

```python
def ups(k):
    """
    Return tally and result functions that compute whether N has exactly K increases.
    """
    # finish the function:
    def f(left, right):
        pass
    pass
```

```python
f, g = ups(3)
process(1200849, f, g)    # Exactly 3 increases: 1 -> 2, 0 -> 8, 4 -> 9
# expected to return True
```

```python
process(94004, f, g)      # 1 increase: 0 -> 4
# expected to return   False 
```

```python
process(122333445, f, g)  # 4 increases: 1 -> 2, 2 -> 3, 3 -> 4, 4 -> 5
# expected to return   False
```

```python
process(0, f, g)          # 0 increases
# expected to return    False
```

```python

```

### P2.(4) [bonus 10 pts]

Implement `at_most`, which returns `True` if the number of increases in a non-negative integer `n` is less than or equal to `k`, and `False` otherwise. Assume `ups` is implemented correctly. You may use any of the functions from previous parts of this question: `sign`, `ramp`, `over_under`, and `process`.

```python
def at_most(n, k):
    """
    Return whether non-negative integer N has at most K increases.
    """
    # finish the function:
    pass
```

```python
at_most(567, 3)
# expected to return True
```

```python
at_most(567, 2)
# expected to return True
```

```python
at_most(567, 1)
# expected to return True
```

```python

```

Tests for p1 (points: 10)
```
>>> assert again(parabola) == 5, 'Test case 1 failed'
```
```
>>> assert again(vee) == 3, 'Test case 2 failed'
```
```
>>> def constant(x):
...     return 42
>>> assert again(constant) == 1, 'Test case 3 failed'
```
```
>>> def periodic(x):
...     return x % 3
>>> assert again(periodic) == 3, 'Test case 5 failed'
```
```
>>> def quadratic(x):
...     return (x - 5) * (x - 10)
>>> assert again(quadratic) == 8, 'Test case 6 failed'
```
```
>>> def abs_diff(x):
...     return abs(x - 4)
>>> assert again(abs_diff) == 5, 'Test case 7 failed'
```
```
>>> def step(x):
...     return x // 2
>>> assert again(step) == 1, 'Test case 8 failed'
```


Tests for p2.(1) (points: 10)
```
>>> assert ramp(123) == True, 'Test case 1 failed'
```
```
>>> assert ramp(1315) == True, 'Test case 2 failed'
```
```
>>> assert ramp(176) == False, 'Test case 3 failed'
```
```
>>> assert ramp(5) == False, 'Test case 4 failed'
```
```
>>> assert ramp(1234) == True, 'Test case 5 failed'
```
```
>>> assert ramp(4321) == False, 'Test case 6 failed'
```
```
>>> assert ramp(1111) == False, 'Test case 7 failed'
```
```
>>> assert ramp(1223) == True, 'Test case 9 failed'
```
```
>>> assert ramp(32123) == False, 'Test case 10 failed'
```


Tests for p2.(2) (points: 10)
```
>>> assert over_under(5)(3) == -1, 'Test case 1 failed'
```
```
>>> assert over_under(5)(5) == 0, 'Test case 2 failed'
```
```
>>> assert over_under(5)(7) == 1, 'Test case 3 failed'
```
```
>>> assert over_under(10)(8) == -1, 'Test case 4 failed'
```
```
>>> assert over_under(10)(10) == 0, 'Test case 5 failed'
```
```
>>> assert over_under(10)(12) == 1, 'Test case 6 failed'
```
```
>>> assert over_under(0)(-1) == -1, 'Test case 7 failed'
```
```
>>> assert over_under(0)(0) == 0, 'Test case 8 failed'
```
```
>>> assert over_under(0)(1) == 1, 'Test case 9 failed'
```


Tests for p2.(3) (points: 10)
```
>>> (f, g) = ups(3)
>>> assert process(1200849, f, g) == False, 'Test case 1 failed'
>>> assert process(94004, f, g) == False, 'Test case 2 failed'
>>> assert process(122333445, f, g) == False, 'Test case 3 failed'
>>> assert process(0, f, g) == False, 'Test case 4 failed'
```
```
>>> (f, g) = ups(2)
>>> assert process(123, f, g) == False, 'Test case 5 failed'
>>> assert process(321, f, g) == True, 'Test case 6 failed'
>>> assert process(111, f, g) == False, 'Test case 7 failed'
>>> assert process(121, f, g) == False, 'Test case 8 failed'
```
```
>>> (f, g) = ups(1)
>>> assert process(10, f, g) == True, 'Test case 9 failed'
>>> assert process(12, f, g) == False, 'Test case 10 failed'
>>> assert process(21, f, g) == True, 'Test case 11 failed'
```


Tests for p2.(4) (points: 10)
```
>>> assert at_most(567, 3) == True, 'Test case 1 failed'
>>> assert at_most(567, 2) == True, 'Test case 2 failed'
>>> assert at_most(567, 1) == True, 'Test case 3 failed'
```
```
>>> assert at_most(1234, 3) == True, 'Test case 4 failed'
```
```
>>> assert at_most(1234, 2) == True, 'Test case 5 failed'
```
```
>>> assert at_most(1111, 0) == True, 'Test case 6 failed'
```
```
>>> assert at_most(1212, 1) == True, 'Test case 7 failed'
```
```
>>> assert at_most(54321, 0) == False, 'Test case 8 failed'
```
```
>>> assert at_most(12321, 2) == True, 'Test case 9 failed'
```
```
>>> assert at_most(12321, 1) == False, 'Test case 10 failed'
```
