# CSCI-A201/597 Introduction to Programming I: 
# Midterm Exam: Programming Part

* Answer each questions as instructed. 

* When you finish the exam:
    * Before submit the file to Autograder: submit the print-out to intructor/TA.
    * Download the notebook file (`.ipynb`) to this link: https://iu.instructure.com/courses/2249164/assignments/16824692. 

```python
from IPython.display import display, HTML
display(HTML("""
<button type="button" id="button_for_indexeddb">Clear JupyterLite local storage</button>
<script>
window.button_for_indexeddb.onclick = function(e) {
    window.indexedDB.open('JupyterLite Storage').onsuccess = function(e) {
        // There are also other tables that I'm not clearing:
        // "counters", "settings", "local-storage-detect-blob-support"
        let tables = ["checkpoints", "files"];

        let db = e.target.result;
        let t = db.transaction(tables, "readwrite");

        function clearTable(tablename) {
            let st = t.objectStore(tablename);
            st.count().onsuccess = function(e) {
                console.log("Deleting " + e.target.result + " entries from " + tablename + "...");
                st.clear().onsuccess = function(e) {
                    console.log(tablename + " is cleared!");
                }
            }
        }

        for (let tablename of tables) {
            clearTable(tablename);
        }
    }
};
</script>
"""))
```

Output:
```
<IPython.core.display.HTML object>
```

## p1.(1).  [10 pts]

```python
# finish the program: 
p1_1 = None
```

```python

```

```python

```

## p1.(2).  [10 pts]

```python
# finish the function:

def p1_2(a):
    pass
```

```python

```

```python

```

## p1.(3).  [10 pts]

```python
# finish the function:

def p1_3(a, b, c):
    pass
```

```python

```

```python

```

## p1.(4)  [10 pts]

```python
# finish the function:

def p1_4(a, b):
    pass
```

```python

```

```python

```

```python

```

### P2 [20 pts]

```python
# P2 (1)


# write your code here

```

```python

```

```python
# P2 (2)


# write your code here

```

```python

```

## p3  [10 pts]

```python
# finish the function:


# write your code here

```

```python

```

```python

```

## p4  [10 pts]

```python
# finish the function:


# write your code here

```

```python

```

```python

```

## p5  [10 pts]

```python
# finish the function:


# write your code here

```

```python

```

```python

```

Tests for p1.(1) (points: 10)
```
>>> def check_p1_1():
...     assert isinstance(p1_1, list), 'not a list'
...     assert len(p1_1) == 96
...     for i in range(5, 101):
...         assert i in p1_1, '{i} should be included'
>>> check_p1_1()
```


Tests for p1.(2) (points: 10)
```
>>> assert p1_2('AEIOU') == 5, 'Test Case 1 Failed'
```
```
>>> assert p1_2('HELLO WORLD') == 3, 'Test Case 2 Failed'
```
```
>>> assert p1_2('BCDFG') == 0, 'Test Case 3 Failed'
```
```
>>> assert p1_2('') == 0, 'Test Case 4 Failed'
```
```
>>> assert p1_2('HEllo') == 1, 'Test Case 5 Failed'
```
```
>>> assert p1_2('happy') == 0, 'Test Case 6 Failed'
```
```
>>> assert p1_2('SUNNY') == 1, 'Test Case 7 Failed'
```
```
>>> assert p1_2('123') == 0, 'Test Case 8 Failed'
```
```
>>> assert p1_2('BCDFGHJKLMNPQRSTVWXYZ' * 100) == 0, 'Test Case 9 Failed'
```


Tests for p1.(3) (points: 10)
```
>>> assert p1_3(22, 25, 28) == 0, 'Test Case 1 Failed'
```
```
>>> assert p1_3(20, 29, 25) == 0
```
```
>>> assert p1_3(10, 15, 30) == 55, 'Test Case 3 Failed'
```
```
>>> assert p1_3(20, 21, 29) == 0, 'Test Case 4 Failed'
```
```
>>> assert p1_3(20, 12, 18) == 30, 'Test Case 5 Failed'
```
```
>>> assert p1_3(20, 25, 50) == 50, 'Test Case 6 Failed'
```
```
>>> assert p1_3(5, 15, 35) == 55, 'Test Case 7 Failed'
```


Tests for p1.(4) (points: 10)
```
>>> assert p1_4('abc', 'abc') == 2, f"Failed Test Case 1: p1_4('abc', 'abc') should return 2"
```
```
>>> assert p1_4('abc', 'xyz') == 0, f"Failed Test Case 2: p1_4('abc', 'xyz') should return 0"
```
```
>>> assert p1_4('aabbcc', 'abbcdd') == 0, f"Failed Test Case 3: p1_4('aabbcc', 'abbcdd') should return 0"
```
```
>>> assert p1_4('abc', 'ab') == 1, f"Failed Test Case 4: p1_4('abc', 'ab') should return 1"
```
```
>>> assert p1_4('', 'xyz') == 0, f"Failed Test Case 5: p1_4('', 'xyz') should return 0"
```
```
>>> assert p1_4('abcdefgh', 'ijklmnop') == 0, f"Failed Test Case 6: p1_4('abcdefgh', 'ijklmnop') should return 0"
```
```
>>> assert p1_4('abcde', 'abcfg') == 2, f"Failed Test Case 7: p1_4('abcde', 'abcfg') should return 2"
```
```
>>> assert p1_4('xyzabc', 'pqrabc') == 2, f"Failed Test Case 8: p1_4('xyzabc', 'pqrabc') should return 2"
```
```
>>> assert p1_4('abcdef', '123def') == 2, f"Failed Test Case 9: p1_4('abcdef', '123def') should return 2"
```
```
>>> assert p1_4('aaaaaa', 'aaaa') == 3, f"Failed Test Case 10: p1_4('aaaaaa', 'aaaa') should return 3"
```


Tests for p2 (points: 20)
```
>>> initial = 100.0
>>> final = 100.0
>>> assert inflation_rate(initial, final) == 0.0, f'Expected 0.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 100.0
>>> final = 120.0
>>> assert inflation_rate(initial, final) == 20.0, f'Expected 20.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 100.0
>>> final = 80.0
>>> assert inflation_rate(initial, final) == -20.0, f'Expected -20.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 50.0
>>> final = 150.0
>>> assert inflation_rate(initial, final) == 200.0, f'Expected 200.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 99.99
>>> final = 100.01
>>> assert inflation_rate(initial, final) == 0.02, f'Expected 0.02, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 1000.0
>>> final = 100.0
>>> assert inflation_rate(initial, final) == -90.0, f'Expected -90.0, but got {inflation_rate(initial, final)}'
```
```
>>> initial = 1000.0
>>> final = 999.99
>>> assert inflation_rate(initial, final) == 0.0
```
```
>>> my_data = [[23.5, 24.0], [24.0, 24.25], [24.25, 24.38], [24.38, 24.75]]
>>> assert average_inflation_rate(my_data) == 1.31
```
```
>>> my_data = [[23.5, 24.0], [24.0, 24.25], [24.25, 24.38]]
>>> assert average_inflation_rate(my_data) == 1.24
```
```
>>> my_data = [[25.0, 25.0], [24.0, 24.0], [30.0, 30.0], [27.0, 27.0], [50.0, 50.0]]
>>> assert average_inflation_rate(my_data) == 0.0, 'Test Case 2 Failed'
```
```
>>> my_data = [[30.0, 28.0], [25.0, 24.0], [40.0, 35.0], [50.0, 55.0], [60.0, 70.0], [80.0, 90.0]]
>>> assert average_inflation_rate(my_data) == 2.67, 'Test Case 3 Failed'
```
```
>>> my_data = [[50.0, 55.0], [55.0, 52.0], [52.0, 60.0], [60.0, 58.0], [58.0, 62.0]]
>>> assert average_inflation_rate(my_data) == 4.7, 'Test Case 5 Failed'
```
```
>>> my_data = [[1000.0, 1200.0], [1200.0, 1000.0], [1000.0, 1500.0], [1500.0, 1300.0], [1300.0, 1700.0], [1700.0, 1600.0], [1600.0, 2000.0]]
>>> assert average_inflation_rate(my_data) == 12.84, 'Test Case 6 Failed'
```
```
>>> my_data = [[1000.0, 1000.01], [1000.01, 1000.02], [1000.02, 1000.03], [1000.03, 1000.04], [1000.04, 1000.05]]
>>> assert average_inflation_rate(my_data) == 0.0, 'Edge Case Failed'
```


Tests for p3 (points: 10)
```
>>> assert attendance_award('PPALLP') == True
```
```
>>> assert attendance_award('PPALLL') == False
```
```
>>> assert attendance_award('PPAALP') == False
```
```
>>> assert attendance_award('PPPPP') == True
```
```
>>> assert attendance_award('PPALP') == True
```
```
>>> assert attendance_award('PPLLL') == False
```
```
>>> assert attendance_award('PLLAP') == True
```
```
>>> assert attendance_award('LLLPP') == False
```
```
>>> assert attendance_award('PLPPLP') == True
```
```
>>> assert attendance_award('PLPLPLP') == True
```
```
>>> assert attendance_award('AAAAA') == False
```
```
>>> assert attendance_award('PPAAAPPP') == False
```
```
>>> assert attendance_award('PPPPPPPP') == True
```
```
>>> assert attendance_award('PPLP') == True
```


Tests for p4 (points: 10)
```
>>> assert maxkd(1234, 1) == 4
```
```
>>> assert maxkd(32749, 2) == 79
```
```
>>> assert maxkd(1917, 2) == 97
```
```
>>> assert maxkd(32749, 18) == 32749
```
```
>>> assert maxkd(12345, 3) == 345
```
```
>>> assert maxkd(98765, 5) == 98765
```
```
>>> assert maxkd(123, 5) == 123
```
```
>>> assert maxkd(112233, 4) == 2233
```
```
>>> assert maxkd(5, 1) == 5
```
```
>>> assert maxkd(5, 0) == 0
```
```
>>> assert maxkd(1111, 2) == 11
```
```
>>> assert maxkd(0, 3) == 0
```
```
>>> assert maxkd(9876543210, 5) == 98765
```
```
>>> assert maxkd(13579, 1) == 9
```
```
>>> assert maxkd(4, 10) == 4
```
```
>>> assert maxkd(54321, 2) == 54
```
```
>>> assert maxkd(32957, 3) == 957
```
```
>>> assert maxkd(888888888, 4) == 8888
```


Tests for p5 (points: 10)
```
>>> original_matrix_1 = [[1, 2, 3], [4, 5, 6], [7, 8, 9]]
>>> expected_output_1 = [[1, 4, 7], [2, 5, 8], [3, 6, 9]]
>>> assert transpose(original_matrix_1) == expected_output_1
```
```
>>> original_matrix_2 = [[1, 2], [3, 4], [5, 6]]
>>> expected_output_2 = [[1, 3, 5], [2, 4, 6]]
>>> assert transpose(original_matrix_2) == expected_output_2
```
```
>>> original_matrix_3 = [[1, 2, 3], [4, 5, 6]]
>>> expected_output_3 = [[1, 4], [2, 5], [3, 6]]
>>> assert transpose(original_matrix_3) == expected_output_3
```
```
>>> original_matrix_4 = [[42]]
>>> expected_output_4 = [[42]]
>>> assert transpose(original_matrix_4) == expected_output_4
```
```
>>> original_matrix_5 = []
>>> expected_output_5 = []
>>> assert transpose(original_matrix_5) == expected_output_5
```
```
>>> original_matrix_6 = [[1, 2, 3]]
>>> expected_output_6 = [[1], [2], [3]]
>>> assert transpose(original_matrix_6) == expected_output_6
```
```
>>> original_matrix_7 = [[1], [2], [3]]
>>> expected_output_7 = [[1, 2, 3]]
>>> assert transpose(original_matrix_7) == expected_output_7
```
